Enhancing Bone Regeneration and Osseointegration using rhPTH(1-34) and Dimeric R25CPTH(1-34) in an Osteoporotic Beagle Model

Author response:

The following is the authors’ response to the original reviews.

Response to Reviewer 1

(Cys25)PTH(1-84) does not show efficacy surpassing that of the previously used rhPTH(1-34). This needs to be discussed biologically and clinically.

Thank you very much for your valuable comments for enhancing the manuscript. We appreciate your input and have noted that this aspect was not addressed in the discussion. The authors have included the following paragraph in discussion section.

“This biological difference is thought to be due to dimeric R25CPTH(1-34) exhibiting a more preferential binding affinity for the RG versus R0 PTH1R conformation, despite having a diminished affinity for either conformation. Additionally, the potency of cAMP production in cells was lower for dimeric R25CPTH compared to monomeric R25CPTH, consistent with its lower PTH1R-binding affinity. (Noh et al., 2024) One of the potential clinical advantages of dimeric R25CPTH(1-34) is its partial agonistic effect in pharmacodynamics. This property may allow for a more fine-tuned regulation of bone metabolism, potentially reducing the risk of adverse effects associated with full agonism, such as hypercalcemia and bone resorption by osteolcast activity. Moreover, the dimeric form may offer a more sustained anabolic response, which could be beneficial in the context of long-term treatment strategies. (Noh et al., 2024) Also, the effects of dimer were prominent, as we mentioned better bone formation than the control group.” (2nd paragraph, Discussion section)

The terms (Cys25)PTH(1-84) and Dimeric R25CPTH(1-34) are being used interchangeably and incorrectly. A unification of these terms is necessary.

We totally agree with the reviewer’s notion. R25CPTH(1-84) represents mutated human PTH, rhPTH(1-34) and dimeric R25CPTH(1-34) are synthesized PTH analogs. To clarified the terminology, we thus have changeed the terminology in the manuscript appear in red.

The figure legend is incorrect. Not all figures are described, and even though there are figures from A to I, only up to E is explained, or the content is different.

We apologize for our negligence. As suggested by a reviewer, we've fixed the figure legends throughout before the list of references in the manuscript as follows.

“Figure legends

Figure 1. Micro-CT analysis (A-D) Experimental design for the controlled delivery of rhPTH(1-34) and dimeric R25CPTH(1-34) in ovariectomized beagle model. Representative images for injection and placement of titanium implant. (E) Micro-CT analysis. bone mineral density (BMD), bone volume (TV; mm3), trabecular number (Tb.N; 1/mm), trabecular thickness (Tb. Th; um), trabecular separation (Tb.sp; ㎛). Error bars indicate standard deviation. Data are shown as mean ± s.d. *p<0.05, **p<0.01, ***p<0.001, n.s., not significant. P, posterior. R, right

Figure 2. (A-I) Histological analysis of the different groups stained in Goldner’s trichrome. The presence of bone is marked by the green color and soft tissue in red. Red arrows indicate the position with soft tissues without bone around the implant threads. The area of bone formed was the widest in the rhPTH(1-34)-treated group. In the dimeric R25CPTH(1-34)treated group, there is a greater amount of bone than vehicle-treated group. Green arrows represent the bone formed over the implant. blue dotted line, margin of bone and soft tissue; Scale bars: 1mm

Figure 3. Histological analysis using Masson trichrome staining results in the rhPTH(1-34) and dimeric R25CPTH(1-34)-treated group (A-L) Masson trichrome-stained sections of cancellous bone in the mandibular bone. The formed bone is marked by the color red. Collagen is stained blue. Black dotted box magnification region of trabecular bone in the mandible. Scale bars, A-C, G-I: 1mm; D-F, J-L: 200 ㎛

Figure 4. Immunohistochemical analysis using TRAP staining for bone remodeling activity (A-L) TRAP staining is used to evaluate bone remodeling by staining osteoclasts. Osteoclasts is presented by the purple color. Black dotted box magnification region of trabecular bone in the mandible. (M, N) The number of TRAP-positive cells in the mandible of the rhPTH(1-34) and dimeric R25CPTH(1-34)-treated beagles. Scale bars, A-C, G-I: 1mm; D-F, J-L: 200 ㎛. Error bars indicate standard deviation. Data are shown as mean ± s.d. *p<0.05, **p<0.01, n.s., not significant

Figure 5. Measurement of biochemical Marker Dynamics in serum. The serum levels of calcium, phosphorus, P1NP, and CTX across three time points (T0, T1, T2) following treatment with dimeric dimeric R25CPTH(1-34), rhPTH(1-34), or control. (A-B) Calcium and phosphorus levels exhibit an upward trend in response to both PTH treatments compared to control, suggesting enhanced bone mineralization. (C) P1NP levels, indicative of bone formation, remain relatively unchanged across time and treatments. (D) CTX levels, associated with bone resorption, show no significant differences between groups. Data points for the dimeric R25CPTH(1-34), rhPTH(1-34), and control are marked by squares, circles, and triangles, respectively, with error bars representing confidence intervals.

Supplementary Figure. Three-dimensional reconstructed image of the bone surrounding the implants. Three-dimensional reconstructed images of the peri-implant bone depicting the osseointegration after different therapeutic interventions. (A) Represents the bone response to recombinant human parathyroid hormone fragment (rhPTH 1-34) treatment, showing the most robust degree of bone formation around the implant in the three groups. (B) Shows the bone response to a modified PTH fragment (dimeric R25CPTH(1-34)), indicating a similar level of bone growth and integration as seen with rhPTH(1-34), although to a slightly lesser extent. (C) Serves as the control group, demonstrating the least amount of bone formation and osseointegration. The upper panel provides a top view of the bone-implant interface, while the lower panel offers a cross-sectional view highlighting the extent of bony ingrowth and integration with the implant surface.”

In Figure 5, although the descriptions of T0, T1, T2 are mentioned in the method section, it would be more clear if there was a timeline like in Figure 1.

Based on the reviewer’s advice, we have indicated the timing of T0, T1, and T2 in the materials & methods section describing the serum biochemical assay, and we have shown a timeline in figure 5.

In Figure 5, instead of having calcium, phosphorus, P1NP, CTX graphs all under Figure 5, it would be more convenient for referencing in the text to label them as Figure 5A, Figure 5B, Figure 5C, Figure 5D.

We totally understood the reviewer’s comment. As the reviewer’s suggested, we have corrected the labeling in the text for figure 5 as follows.

“The levels of calcium, phosphorus, CTX, and P1NP were analyzed over time using RM-ANOVA (Figure 5). There were no significant differences between the groups for calcium and phosphorus at time points T0 and T1 (Figure 5A). However, after the PTH analog was administered at T2 (Figure 5A), the levels were highest in the rhPTH(1-34) group, followed by the dimeric R25CPTH(1-34) group, and then, lowest in the control group, which was statistically significant (Figure 5B,C). (P < 0.05) The differences between the groups over time for CTX and P1NP were not statistically significant (Figure 5D, E).”

Significance should be indicated in the figure (no asterisk present).

As the reviewer’s comment, we put the asterisk in the figure 5.

Addition of Figures in Text:

Line 112: change from "figure 2" to "figure 1" / Line 115: mention "figure 1. E"

Line 120: refer to "figure 1. E" / Line 123: change from "figure 3" to "figure 2"

Line 128: refer to "figure 2.A-C" / Line 137: mention "figure 3"

Line 138: refer to "figure 3. A-L" / Line 143: mention "figure 3. A-L"

Line 144: refer to "figure 3. E,F,K,L" / Line 148: mention "figure 4"

Line 150: refer to "figure 4 M,N" / Line 152: mention "figure 4. M,N"

Line 155: refer to "figure 5" / Line 157: mention "figure 5"

Line 159: refer to "figure 5" / Line 171: mention "figure 1 E"

Line 175: refer to "figure 2 M, N"/ Line 194: mention "figure 3"

Above all, thank you for the reviewer’s notion. We corrected detailed figure labeling in text to red color.

Response to Reviewer 2

First, the authors should clarify why they compared the effects of rhPTH(1-34) and of dimeric R25C2 PTH(1-34)? In most of the parameters, rhPTH(1-34) seems to be superior to dimeric R25C2 PTH(1-34). Why did the authors insist that the anabolic effects of dimer were prominent? Even though implication of dimeric R25C2 PTH(1-34) was drawn from genetic mutation studies, the authors should describe more clearly in the discussion the potential clinical benefits of the dimeric R25C2 PTH(1-34) compared to rhPTH(1-34), especially if dimeric R25C2 PTH(1-34) has just partial agonistic effect in pharmacodynamics.

Thank you for your insightful comments and questions regarding our results between rhPTH(1-34) and dimeric R25CPTH(1-34). rhPTH(1-34) is a well-characterized therapy for osteoporosis. In this study, rhPTH(1-34) generally showed superior outcomes in most parameters tested, the dimeric R25CPTH(1-34) exhibited specific anabolic effects that are not as pronounced with rhPTH(1-34). We recognized R25CPTH(1-34) as a anabolic effector. One of the potential advantages of dimeric R25CPTH(1-34) is its partial agonistic effect in pharmacodynamics. This property may allow for a more fine-tuned regulation of bone metabolism, potentially reducing the risk of adverse effects associated with full agonism, such as hypercalcemia and bone resorption by osteolast activity. Moreover, the dimeric form may offer a more sustained anabolic response, which could be beneficial in the context of long-term treatment strategies. Also, based on our results, we notes that the effects of dimer were prominent, as we mentioned better bone formation than the control group. We appreciate your input and have noted that this aspect was not addressed in the discussion. As a result, we have included the following paragraph in discussion section.

“This biological difference is thought to be due to dimeric R25CPTH(1-34) exhibiting a more preferential binding affinity for the RG versus R0 PTH1R conformation, despite having a diminished affinity for either conformation. Additionally, the potency of cAMP production in cells was lower for dimeric R25CPTH compared to monomeric R25CPTH, consistent with its lower PTH1R-binding affinity. (Noh et al., 2024) One of the potential clinical advantages of dimeric R25CPTH(1-34) is its partial agonistic effect in pharmacodynamics. This property may allow for a more fine-tuned regulation of bone metabolism, potentially reducing the risk of adverse effects associated with full agonism, such as hypercalcemia and bone resorption by osteolcast activity. Moreover, the dimeric form may offer a more sustained anabolic response, which could be beneficial in the context of long-term treatment strategies. (Noh et al., 2024) Also, the effects of dimer were prominent, as we mentioned better bone formation than the control group.” (2nd paragraph, Discussion section)

Second, please describe the intermittent and continuous application of PTH analogues. Many of the readers may misunderstand that the authors' daily injection of PTHs were actually to mimic the clinical intermittent application or continuous one. Incorporation of the author's intention for experimental design would be more helpful for readers.

Thank you for your insightful comments regarding the need for clearer differentiation between intermittent and continuous applications of PTH analogs in this study. We appreciate your concern that the readers may not fully grasp whether our daily injection protocol was intended to mimic clinical intermittent or continuous PTH administration. To address this, we have revised the manuscript to explicitly clarify that the daily injections of rhPTH(1-34) and dimeric R25CPTH(1-34) were designed to simulate the intermittent dosing regimen commonly used in clinical practice. This regimen is known to maximize the anabolic effects on bone while minimizing potential catabolic actions associated with more frequent or continuous hormone exposure. We have added detailed explanations in the Introduction, Methods, and Discussion sections to help readers understand our experimental design and its relevance to clinical settings.

Introduction section

“Administration of prathyroid hormone (PTH) analogs can be categorized into two distinct protocols: intermittent and continuous. Intermittent rhPTH(1-34) therapy, typically characterized by daily injections, is clinically used to enhance bone formation and strength. This method leverages the anabolic effects of rhPTH(1-34) without significant bone resorption, which can occur with more frequent or continuous exposure. On the other hand, continuous rhPTH(1-34) exposure, often modeled in research as constant infusion, tends to accelerate bone resorption activities, potentially leading to bone loss (Silva and Bilezikian, 2015; Jilka, 2007). Understanding these differences is crucial for interpreting the therapeutic implications of rhPTH(1-34) in bone health.”

Silva, B. C., & Bilezikian, J. P. (2015). Parathyroid hormone: anabolic and catabolic actions on the skeleton. Current Opinion in Pharmacology, 22, 41-50.

Jilka, R. L. (2007). Molecular and cellular mechanisms of the anabolic effect of intermittent PTH. Bone, 40(6), 1434-1446.

Materials and Methods section

“Each animal received one injection per day, aimed at replicating the intermittent rhPTH(1-34) exposure proven beneficial for bone regeneration and overall skeletal health in clinical settings (Neer et al., 2001; Kendler et al., 2018). This regimen was chosen to investigate the potential anabolic effects of these specific PTH analogs under conditions closely resembling therapeutic use.”

Neer, R. M., Arnaud, C. D., Zanchetta, J. R., Prince, R., Gaich, G. A., Reginster, J. Y., Hodsman, A. B., Eriksen, E. F., Ish-Shalom, S., Genant, H. K., Wang, O., and Mitlak, B. H. (2001). Effect of Parathyroid Hormone (1-34) on Fractures and Bone Mineral Density in Postmenopausal Women with Osteoporosis. The New England Journal of Medicine, 344(19), 1434-1441.

Kendler, D. L., Marin, F., Zerbini, C. A. F., Russo, L. A., Greenspan, S. L., Zikan, V., Bagur, A., Malouf-Sierra, J., Lakatos, P., Fahrleitner-Pammer, A., Lespessailles, E., Minisola, S., Body, J. J., Geusens, P., Moricke, R., & Lopez-Romero, P. (2018). Effects of Teriparatide and Risedronate on New Fractures in Post-Menopausal Women with Severe Osteoporosis (VERO): A Multicenter, Double-Blind, Double-Dummy, Randomized Controlled Trial. The Lancet, 391(10117), 230-240.

Discussion section

“The use of daily injections in this study was intended to simulate intermittent PTH therapy, a well-established clinical approach for managing osteoporosis and enhancing bone regeneration. Intermittent administration of PTH, as opposed to continuous exposure, is critical for maximizing the anabolic response while minimizing the catabolic effects that are associated with higher frequency or continuous hormone levels. Our findings support the notion that even with daily administration, both rhPTH(1-34) and dimeric dimeric R25CPTH(1-34) promote bone formation and osseointegration, consistent with the outcomes expected from intermittent therapy. It’s important for future research to consider the dosage and timing of administration to further optimize the therapeutic benefits of PTH analogs (Dempster et al., 2001; Hodsman et al., 2005).”

Dempster, D. W., Cosman, F., Kurland, E. S., Zhou, H., Nieves, J., Woelfert, L., Shane, E., Plavetic, K., Müller, R., Bilezikian, J., & Lindsay, R. (2001). Effects of Daily Treatment with Parathyroid Hormone on Bone Microarchitecture and Turnover in Patients with Osteoporosis: A Paired Biopsy Study. Journal of Bone and Mineral Research, 16(10), 1846-1853.

Hodsman, A. B., Bauer, D. C., Dempster, D. W., Dian, L., Hanley, D. A., Harris, S. T., Kendler, D. L., McClung, M. R., Miller, P. D., Olszynski, W. P., Orwoll, E., Yuen, C. K. (2005). Parathyroid Hormone and Teriparatide for the Treatment of Osteoporosis: A Review of the Evidence and Suggested Guidelines for Its Use. Endocrine Reviews, 26(5), 688-703.

Third, please unify the nomenclature. Ensure consistency in the nomenclature throughout the article. Unify the naming conventions for PTH analogues, such as rhPTH(1-34) vs teriparatide and (Cys25)PTH(1-84) vs R25CPTH(1-34) vs R25CPTH(1-34) vs (1-84). Choose one nomenclature for each analogue and use it consistently throughout the article.

We totally agree with the reviewer’s notion. R25CPTH(1-84) represents mutated human PTH, rhPTH(1-34) and dimeric R25CPTH(1-34) are synthesized PTH analogs. To clarified the terminology, we thus have changed the terminology in the manuscript appear in red.

Response to Reviewer 3

I would recommend to rewrite the manuscript in a form that is more understandable to the readers. In fact, it appears to me that this work was originally formatted in a way that would need the Materials and Methods to precede the results. As presented (and as requested by the eLife formatting) the Materials and Methods are available only at the end of the reading and, as a consequence, the readers needs to refer to the Materials and Methods to have a general and initial understanding of the study design (i.e. type of treatment for each group, etc are not well specified in the Results section).

Thank you for you constructive comments and suggestions regarding the manuscript. We appreciate your feedback on the organization of the manuscript entirely. As reviewer mentioned, Materials and methods were placed after the discussion section in accordance with the format of the elife journal. For a better and initial understanding, a description of each experimental group has been added to the Results section as follow. Thank you again for your valuable comments.

“To investigate evaluating and comparing the efficacy of rhPTH(1-34) and the dimeric R25CPTH(1-34) in promoting bone regeneration and healing in a clinically relevant animal model. In our study, beagle dogs were selected as the model due to their anatomical similarity to human oral structures, suitable size for surgeries, human-like bone turnover rates, and established oral health profiles, ensuring comparable and ethically sound research outcomes. The normal saline injected-control group, injected with 40ug/day PTH (Forsteo, Eli Lilly) group, and 40ug/day PTH analog-injected group. Animals in each group were injected subcutaneously for 10 weeks.”

Changes from the earlier version to the revised version

Abstract

Abstract

Introduction

Despite continuous parathyroid hormone (PTH) exposure potentially accelerating bone resorption, intermittent PTH administration has shown anabolic effects on bone microarchitecture. This study investigates the therapeutic impacts of two PTH analogs, rhPTH(1-34) and dimeric R25CPTH(1-34) on bone regeneration and osseointegration in a postmenopausal osteoporosis animal model.

Methods

Twelve female beagles, osteoporotic post-ovariectomy, underwent implant surgeries. Animals were divided into three groups: control, daily rhPTH(1-34) injection, and daily dimeric R25CPTH(1-34) injection. After 10 weeks, bone regeneration and implant osseointegration were evaluated using micro-CT, histological/histomorphometric analyses, and serum biochemical analysis.

Results

Earlier version: While the rhPTH(1-34) group demonstrated improved microarchitectural characteristics, such as BMD, BV, trabecular metrics, and osseointegration degree, the dimeric R25CPTH(1-34) group showed similarly enhanced anabolic effects around the titanium implants, albeit these were less pronounced than the rhPTH(1-34), yet significantly better than the control group. Histologic and TRAP analyses revealed that both PTH analogs significantly promoted bone regeneration and remodeling, especially in artificially created bone defects.

Revised version: While the rhPTH(1-34) group demonstrated improved microarchitectural characteristics, such as BMD, BV, trabecular metrics, and osseointegration degree, the dimeric R25CPTH(1-34) group showed similarly enhanced anabolic effects around the titanium implants, albeit these were less pronounced than the rhPTH(1-34), yet significantly better than the control group. Histologic and TRAP assays revealed that both PTH analogs significantly promoted bone regeneration and remodeling, especially in artificially created bone defects.

Conclusion

This study demonstrated the therapeutic effects of rhPTH(1-34) and dimeric R25CPTH(1-34) on bone regeneration and titanium osseointegration in a beagle model with osteoporosis.

Introduction

Earlier version: Since excessive levels of circulating parathyroid hormone (PTH) increase osteoclastic activity and accelerate bone resorption, it might seem paradoxical that PTH can also be used as a treatment modality for diseases with bone loss, such as osteoporosis (Jilka, 2007). Intermittent administration of PTH unlikely continuous exposure showed anabolic effects, indicating different responses relating to bone microarchitecture depending on the dose and frequency (Silva and Bilezikian, 2015). However, the underlying mechanisms remain largely unknown, although it was shown that PTH binds through PTH type 1 receptor (PTH1R) and that G-protein coupled receptors (GPCR) are associated with the protein kinase A (PKA)-dependent pathway, thereby demonstrating primary anabolic action on bone (Cheloha et al., 2015; Jilka, 2007). The anabolic effect of intermittent PTH administration is mediated by the downregulation of the Wnt/beta-catenin signaling pathway, which upregulates the transcriptional expression of growth factors, such as IGF1, FGF2, and Runx2, which are essential for the proliferation and differentiation of osteoblasts, leading to an increased number of osteoblasts and survival (Krishnan et al., 2006; Lee et al., 2009).

Revised version: Bone regeneration is a complex physiological process that is essential for the successful integration of dental implants into the jawbone. Osseointegration, the direct structural and functional connection between living bone and an implant’s surface, is a critical determinant of implant stability and long-term viability in dental reconstructive therapy. Since excessive levels of circulating parathyroid hormone (PTH) increase osteoclastic activity and accelerate bone resorption, it might seem paradoxical that PTH can also be used as a treatment modality for diseases with bone loss, such as osteoporosis (Jilka, 2007). Administration of prathyroid hormone (PTH) analogs can be categorized into two distinct protocols: intermittent and continuous. Intermittent rhPTH(1-34) therapy, typically characterized by daily injections, is clinically used to enhance bone formation and strength. This method leverages the anabolic effects of rhPTH(1-34) without significant bone resorption, which can occur with more frequent or continuous exposure. On the other hand, continuous rhPTH(1-34) exposure, often modeled in research as constant infusion, tends to accelerate bone resorption activities, potentially leading to bone loss (Silva and Bilezikian, 2015; Jilka, 2007). Understanding these differences is crucial for interpreting the therapeutic implications of rhPTH(1-34) in bone health. Intermittent administration of PTH unlikely continuous exposure showed anabolic effects, indicating different responses relating to bone microarchitecture depending on the dose and frequency (Silva and Bilezikian, 2015). However, the underlying mechanisms remain largely unknown, although it was shown that PTH binds through PTH type 1 receptor (PTH1R) and that G-protein coupled receptors (GPCR) are associated with the protein kinase A (PKA)-dependent pathway, thereby demonstrating primary anabolic action on bone (Cheloha et al., 2015; Jilka, 2007). The anabolic effect of intermittent PTH administration is mediated by the downregulation of the Wnt/beta-catenin signaling pathway, which upregulates the transcriptional expression of growth factors, such as IGF1, FGF2, and Runx2, which are essential for the proliferation and differentiation of osteoblasts, leading to an increased number of osteoblasts and survival (Krishnan et al., 2006; Lee et al., 2009).

Earlier version: A PTH analog, teriparatide (rhPTH(1-34)), demonstrated its promising anabolic effects in a fracture prevention trial (FPT) (Neer et al., 2001), leading to its approval by the United States Food and Drugs Administration, making it the first anabolic agent for postmenopausal osteoporotic women. Additionally, its unique anabolic features, which are contrasted with antiresorptive, lead to an increased application of rhPTH(1-34), whereby it was used in both metabolic and pathologic bone diseases alongside various other conditions where bone formation occurred (Uusi-Rasi et al., 2005). Conversely, recent concerns regarding the development of osteonecrosis of the jaw (ONJ) have appeared in association with the use of antiresorptive, meaning rhPTH(1-34) has gained attention for its potential to reduce the risk of ONJ and its therapeutic effects in treating ONJ (Jung et al., 2017; Kakehashi et al., 2015). Although in vivo studies on fracture healing, bone augmentation, and titanium osseointegration effects of rhPTH(1-34) have been attempted, they have generally been limited to pilot studies in rodents (Gomes-Ferreira et al., 2020; Jung et al., 2021; Yu and Su, 2020).

Revised version: A PTH analog, teriparatide [recombinant human PTH(1-34){rhPTH(1-34)}], demonstrated its promising anabolic effects in a fracture prevention trial (FPT) (Neer et al., 2001), leading to its approval by the United States Food and Drugs Administration, making it the first anabolic agent for postmenopausal osteoporotic women. Additionally, its unique anabolic features, which are contrasted with antiresorptive, lead to an increased application of rhPTH(1-34), whereby it was used in both metabolic and pathologic bone diseases alongside various other conditions where bone formation occurred (Uusi-Rasi et al., 2005). Conversely, recent concerns regarding the development of osteonecrosis of the jaw (ONJ) have appeared in association with the use of antiresorptive, meaning rhPTH(1-34) has gained attention for its potential to reduce the risk of ONJ and its therapeutic effects in treating ONJ (Jung et al., 2017; Kakehashi et al., 2015). Although in vivo studies on fracture healing, bone augmentation, and titanium osseointegration effects of rhPTH(1-34) have been attempted, they have generally been limited to pilot studies in rodents (Gomes-Ferreira et al., 2020; Jung et al., 2021; Yu and Su, 2020).

Earlier version: Interestingly, only nine mutations have been discovered since the PTH amino acid and nucleotide sequences were confirmed (Lee et al., 2020; Schipani et al., 1999). Among them, the PTH R25C mutation was discovered in three siblings with familial idiopathic hypoparathyroidism, which consisted of a homozygous arginine to cysteine mutation at residue 25 (R25C) in the mature PTH(1-84) polypeptide and exhibited distinct characteristics from the others (Lee and Lee, 2022). The other mutations located in the prepro-leader region of the hormone resulted in defective synthesis and secretion; however, the PTH R25C mutation is located within the mature bioactive domain of PTH and does not affect synthesis or secretion (Lee and Lee, 2022). Although the capacity of R25CPTH(1-34) to bind to the PTH1R and stimulate cAMP production was slightly lower in the human osteoblast-derived SaOS-2 cell line (Lee and Lee, 2022), yet it showed comparable anabolic activity in a mouse model (Bae et al., 2016). Furthermore, the dimeric formation of the R25CPTH (1-84) peptide, presumably through disulfide bonding of the cysteine residues, has implied that dimeric R25CPTH(1-34) might partake in unique biological actions, which potentiate its clinical applications (Park et al., 2021).

Revised version: Interestingly, only nine mutations have been discovered since the PTH amino acid and nucleotide sequences were confirmed (Lee et al., 2020; Schipani et al., 1999). Among them, the PTH R25C mutation was discovered in three siblings with familial idiopathic hypoparathyroidism, which consisted of a homozygous arginine to cysteine mutation at residue 25 (R25C) in the mature PTH(1-84) polypeptide and exhibited distinct characteristics from the others (Lee and Lee, 2022). The other mutations located in the prepro-leader region of the hormone resulted in defective synthesis and secretion; however, the PTH R25C mutation is located within the mature bioactive domain of PTH and does not affect synthesis or secretion (Lee and Lee, 2022). Although the capacity of R25CPTH(1-34) to bind to the PTH1R and stimulate cAMP production was slightly lower in the human osteoblast-derived SaOS-2 cell line (Lee and Lee, 2022), yet it showed comparable anabolic activity in a mouse model (Bae et al., 2016). Furthermore, the dimeric formation of the R25CPTH(1-34) peptide, presumably through disulfide bonding of the cysteine residues, has implied that dimeric R25CPTH(1-34) might partake in unique biological actions, which potentiate its clinical applications (Park et al., 2021).

Earlier version: In addition to the limited effectiveness toward certain types of fractures, such as non-vertebral fractures, and the disadvantage of its limited duration of use, the administration of rhPTH(1-34) for osteoporosis has raised concerns about its potential to induce cortical porosity, despite showing favorable results in the treatment of trabecular microarchitecture (Lindsay et al., 2016). This has led to concerns regarding the widespread use of PTH. However, recent notable studies have indicated that cortical porosity is not induced when PTH is administered weekly (Mosekilde et al., 1995), which is in contrast to daily administrations. (Yamamoto et al., 2016; Zebaze et al., 2017) These studies suggest that the frequency and dosage of PTH administration can significantly affect the bone response (Hock et al., 1992), while different forms of PTH might also induce different biological responses, which means deeper investigations into the therapeutic effects of PTH are required (Bellido et al., 2005). Therefore, in this study, the authors used a large animal model that mimics postmenopausal osteoporosis to investigate the therapeutic effects of two PTH analogs, rhPTH(1-34) and dimeric R25CPTH(1-34), on bone regeneration and osseointegration.

Revised version: In addition to the limited effectiveness toward certain types of fractures, such as non-vertebral fractures, and the disadvantage of its limited duration of use, the administration of rhPTH(1-34) for osteoporosis has raised concerns about its potential to induce cortical porosity, despite showing favorable results in the treatment of trabecular microarchitecture (Lindsay et al., 2016). This has led to concerns regarding the widespread use of PTH. However, recent notable studies have indicated that cortical porosity is not induced when PTH is administered weekly (Mosekilde et al., 1995), which is in contrast to daily administrations. (Yamamoto et al., 2016; Zebaze et al., 2017) These studies suggest that the frequency and dosage of PTH administration can significantly affect the bone response (Hock et al., 1992), while different forms of PTH might also induce different biological responses, which means deeper investigations into the therapeutic effects of PTH are required (Bellido et al., 2005).

Added in the revised version: Therefore, in this study, the authors used a large animal model that mimics postmenopausal osteoporosis to investigate the therapeutic effects of two PTH analogs, rhPTH(1-34) and dimeric R25CPTH(1-34), on bone regeneration and osseointegration.

Results

Microarchitectural and histological analysis of titanium osseointegration

Added in the revised version: To investigate evaluating and comparing the efficacy of rhPTH(1-34) and the dimeric R25CPTH(1-34) in promoting bone regeneration and healing in a clinically relevant animal model. In our study, beagle dogs were selected as the model due to their anatomical similarity to human oral structures, suitable size for surgeries, human-like bone turnover rates, and established oral health profiles, ensuring comparable and ethically sound research outcomes. The normal saline injected-control group, injected with 40ug/day PTH (Forsteo, Eli Lilly) group, and 40ug/day dimeric R25CPTH(1-34) -injected group. Animals in each group were injected subcutaneously for 10 weeks.

Earlier version: Various characteristics of the right mandible were evaluated in this analysis, including bone mineral density (BMD), bone volume (BV), trabecular number (Tb.N), trabecular thickness (Tb.Th), and trabecular separation (Tb.S). (Figure 2) The group administered PTH(1-34) presented consistently higher values of BMD, BV, Tb.N, and Tb.Th,, and lower value of Tb.S compared to the control and (Cys25)PTH(1-84) groups, indicating that PTH(1-34) administration enhanced titanium osseointegration. Moreover, this was consistent for all three titanium implants, irrespective of artificial bone defects (second implant) or bone grafting (third implant). Interestingly, the (Cys25)PTH(1-84) group showed similar trends for the anabolic effects related to the titanium implants. Morphometric analysis indicated that (Cys25)PTH(1-84) administration enhanced the BMD, BV, and other trabecular indices, resulting in a higher degree of osseointegration than in the control group (P < 0.05), although lower than in the PTH(1-34) group (P > 0.05).

Revised version: Various characteristics of the right mandible were evaluated in this analysis, including bone mineral density (BMD), bone volume (BV), trabecular number (Tb.N), trabecular thickness (Tb.Th), and trabecular separation (Tb.S) (Figure 1). The group administered rhPTH(1-34) presented consistently higher values of BMD, BV, Tb.N, and Tb.Th,, and lower value of Tb.S compared to the control and dimeric R25CPTH(1-34) groups, indicating that rhPTH(1-34) administration enhanced titanium osseointegration (Figure 1E). Moreover, this was consistent for all three titanium implants, irrespective of artificial bone defects (second implant) or bone grafting (third implant). Interestingly, the dimeric R25CPTH(1-34) group showed similar trends for the anabolic effects related to the titanium implants. Morphometric analysis indicated that dimeric R25CPTH(1-34) administration enhanced the BMD, BV, and other trabecular indices, resulting in a higher degree of osseointegration than in the control group (P < 0.05), although lower than in the rhPTH(1-34) group (P > 0.05) (Figure 1E).

Added in the revised version: Micro-CT analysis (A-D) Experimental design for the controlled delivery of rhPTH(1-34) and dimeric R25CPTH(1-34) in ovariectomized beagle model. Representative images for injection and placement of titanium implant. (E) Micro-CT analysis. bone mineral density (BMD), bone volume (TV; mm3), trabecular number (Tb.N; 1/mm), trabecular thickness (Tb. Th; um), trabecular separation (Tb.sp; ㎛). Error bars indicate standard deviation. Data are shown as mean ± s.d. *p<0.05, **p<0.01, ***p<0.001, n.s., not significant. P, posterior. R, right

Removed in the revised version: A) Experimental timeline, B) subcutaneous injection of PTH analogs, C) intraoral photograph, D) micro-CT scan of implant inserted in the right mandible.

Removed in the revised version: Microarchitectural CT analysis.

Removed in the revised version: (A) Bone mineral density (BMD), (B) bone volume (BV; mm3), (C) trabecular number (Tb.N; 1/mm), (D) trabecular thickness (Tb.Th; μm), (E) trabecular separation (Tb.S; um).

Removed in the revised version: * indicates p < 0.05, ** indicates p < 0.01, and *** indicates p < 0.001

Earlier version: Histological analysis further clarified the micro-CT results. The implants in the control group showed physiological bone osseointegration around the titanium implant; however, insufficient bone– implant contact and exposure of implant threads were observed. (Figure 3) This was especially evident on the buccal side of the implant, which was more vulnerable due to the bundle bone structure. Alternatively, the titanium implants in the PTH(1-34) group were in full contact with the green-stained mineralized bone. The (Cys25)PTH(1-84) group showed a better pattern of osseointegration compared to the control group, although the bone–implant contact was lower than in the PTH(1-34) group. Notably, both the PTH(1-34) and (Cys25)PTH(1-84) groups presented evidence of bone regeneration for the second implant, whereby a bone defect was created prior to placing the implant. However, both were insufficient compared to the control group and did not illustrate any bone filling. The measured bone–implant contact ratio was 18.32 ± 16.19% for the control group, 48.13 ± 29.81% for the group, and 39.53 ± 26.17% (P < 0.05).

Revised version: Histological analysis further clarified the micro-CT results. The implants in the control group showed physiological bone osseointegration around the titanium implant; however, insufficient bone– implant contact and exposure of implant threads were observed (Figure 2). This was especially evident on the buccal side of the implant, which was more vulnerable due to the bundle bone structure. Alternatively, the titanium implants in the rhPTH(1-34) group were in full contact with the green-stained mineralized bone. The dimeric R25CPTH(1-34)group showed a better pattern of osseointegration compared to the control group, although the bone–implant contact was lower than in the rhPTH(1-34) group (Figure 2A-C). Notably, both the rhPTH(1-34) and dimeric R25CPTH(1-34) groups presented evidence of bone regeneration for the second implant, whereby a bone defect was created prior to placing the implant. However, both were insufficient compared to the control group and did not illustrate any bone filling. The measured bone–implant contact ratio was 18.32 ± 16.19% for the control group, 48.13 ± 29.81% for the group, and 39.53 ± 26.17% (P < 0.05).

Added in the revised version: (A-I) Histological analysis of the different groups stained in Goldner’s trichrome. The presence of bone is marked by the green color and soft tissue in red. Red arrows indicate the position with soft tissues without bone around the implant threads. The area of bone formed was the widest in the rhPTH(1-34)-treated group. In the dimeric R25CPTH(1-34)treated group, there is a greater amount of bone than vehicle-treated group. Green arrows represent the bone formed over the implant. blue dotted line, margin of bone and soft tissue; Scale bars: 1mm

Removed in the revised version: (1) Histological analysis using Goldner’s trichrome staining for the regenerated bone around the titanium implant (right mandible), (2) histological analysis using Masson trichrome staining for the bone remodeling pattern (left mandible).

Histological and TRAP analyses of bone regeneration

Earlier version: Artificial bilateral bone defects were created in the left mandible and the effects of the two PTH analogs were evaluated on the bone regeneration, with one left unfilled and the other filled with a bone graft. Figure 4 demonstrates the effects of PTH(1-34) and (Cys25)PTH(1-84) on bone regeneration compared to the control group. Following the formation of the bone defects, the PTH(1-34) and (Cys25)PTH(1-84) groups achieved sufficient morphological bone regeneration over a period of 10 weeks, while the control group exhibited morphological incompleteness over the same period. The PTH(1-34) group exhibited a mature trabecular architecture, while the (Cys25)PTH(1-84) group showed a similar morphology, although some immature bone formation remained stained blue in the Masson trichrome staining analysis. While there was no clear difference in the bone defect between the sites with and without bone grafting, the site where bone grafting occurred exhibited a more mature bone morphology, indicating that a xenograft-maintained space for new bone formation with osteoconductive effects.

Revised version: Artificial bilateral bone defects were created in the left mandible and the effects of the two PTH analogs were evaluated on the bone regeneration, with one left unfilled and the other filled with a bone graft. Figure 3 demonstrates the effects of rhPTH(1-34) and dimeric R25CPTH(1-34) on bone regeneration compared to the control group (Figure 3A-L). Following the formation of the bone defects, the rhPTH(1-34) and dimeric R25CPTH(1-34)groups achieved sufficient morphological bone regeneration over a period of 10 weeks, while the control group exhibited morphological incompleteness over the same period. The rhPTH(1-34) group exhibited a mature trabecularcfghnc architecture, while the dimeric R25CPTH(1-34)group showed a similar morphology, although some immature bone formation remained stained blue in the Masson trichrome staining analysis (Figure 3A-L). While there was no clear difference in the bone defect between the sites with and without bone grafting (Figure 3E,F,K,L), the site where bone grafting occurred exhibited a more mature bone morphology, indicating that a xenograft-maintained space for new bone formation with osteoconductive effects.

Added in the revised version: Histological analysis using Masson trichrome staining results in the rhPTH(1-34) and dimeric R25CPTH(1-34)-treated group (A-L) Masson trichrome-stained sections of cancellous bone in the mandibular bone. The formed bone is marked by the color red. Collagen is stained blue. Black dotted box magnification region of trabecular bone in the mandible. Scale bars, A-C, G-I: 1mm; D-F, J-L: 200 ㎛

Removed in the revised version: Immunohistochemical analysis using TRAP staining for bone remodeling activity. (A–L) TRAP-positive cells (M–N) indicates p < 0.05, ** indicates p < 0.01, and *** indicates p < 0.001.

Earlier version: The capability of PTH(1-34) and (Cys25)PTH(1-84) in bone remodeling was evaluated by TRAP immunohistochemical staining. (Figure 5) Both the PTH(1-34) and (Cys25)PTH(1-84) groups showed a significantly higher number of TRAP+ cells at both bone defects, with and without a xenograft, compared to the control group. (P < 0.05) In addition, the number of TRAP+ cells in the (Cys25)PTH(1-84) group was significantly higher than in the vehicle, yet lower than in the PTH(1-34) group.

Revised version: The capability of rhPTH(1-34) and dimeric R25CPTH(1-34) in bone remodeling was evaluated by tartrate-resistant acid phosphatase (TRAP) immunohistochemical staining. (Figure 4) Both the rhPTH(1-34) and dimeric R25CPTH(1-34) groups showed a significantly higher number of TRAP+ cells at both bone defects, with and without a xenograft, compared to the control group (Figure 4M,N). (P < 0.05) In addition, the number of TRAP+ cells in the dimeric R25CPTH(1-34)group was significantly higher than in the vehicle, yet lower than in the rhPTH(1-34) group (Figure 4M,N).

Added in the revised version: Immunohistochemical analysis using TRAP staining for bone remodeling activity (A-L) TRAP staining is used to evaluate bone remodeling by staining osteoclasts. Osteoclasts is presented by the purple color. Black dotted box magnification region of trabecular bone in the mandible. (M, N) The number of TRAP-positive cells in the mandible of the rhPTH(1-34) and dimeric R25CPTH(1-34)-treated beagles. Scale bars, A-C, G-I: 1mm; D-F, J-L: 200 ㎛. Error bars indicate standard deviation. Data are shown as mean ± s.d. *p<0.05, **p<0.01, n.s., not significant

Serum biochemical analysis

Removed in the revised version: Serum biochemical analysis

Earlier version: The levels of calcium, phosphorus, CTX, and P1NP were analyzed over time using RM-ANOVA. There were no significant differences between the groups for calcium and phosphorus at time points T0 and T1. However, after the PTH analog was administered at T2, the levels were highest in the PTH(1-34) group, followed by the (Cys25)PTH(1-84) group, and then, lowest in the control group, which was statistically significant. (P < 0.05) The differences between the groups over time for CTX and P1NP were not statistically significant.

Revised version: The levels of calcium, phosphorus, CTX, and P1NP were analyzed over time using RM-ANOVA (Figure 5). There were no significant differences between the groups for calcium and phosphorus at time points T0 and T1 (Figure 5A). However, after the PTH analog was administered at T2 (Figure 5A), the levels were highest in the rhPTH(1-34) group, followed by the dimeric R25CPTH(1-34) group, and then, lowest in the control group, which was statistically significant (Figure 5B,C). (P < 0.05) The differences between the groups over time for CTX and P1NP were not statistically significant (Figure 5D, E).

Added in the revised version: Measurement of biochemical Marker Dynamics in serum. The serum levels of calcium, phosphorus, P1NP, and CTX across three time points (T0, T1, T2) following treatment with dimeric dimeric R25CPTH(1-34), rhPTH(1-34), or control. (A-B) Calcium and phosphorus levels exhibit an upward trend in response to both PTH treatments compared to control, suggesting enhanced bone mineralization. (C) P1NP levels, indicative of bone formation, remain relatively unchanged across time and treatments. (D) CTX levels, associated with bone resorption, show no significant differences between groups. Data points for the dimeric R25CPTH(1-34), rhPTH(1-34), and control are marked by squares, circles, and triangles, respectively, with error bars representing confidence intervals.

Discussion

Earlier version: This study investigated the therapeutic effects of rhPTH(1-34) and dimeric R25CPTH(1-34) on bone regeneration and osseointegration in a large animal model with postmenopausal osteoporosis. rhPTH(1-34) and dimeric R25CPTH(1-34) have shown significant clinical efficacy, and although there have been a few studies investigating their effects on bone regeneration in rodents (Garcia et al., 2013), the authors in this study aimed to investigate the effects using a large animal model that more accurately mimics osteoporotic humans (Cortet, 2011). In the evaluation of titanium osseointegration, the PTH(1-34) group consistently exhibited enhanced bone mineral density (BMD), bone volume (BV), and other key parameters, thereby indicating superior titanium osseointegration compared to the control and dimeric R25CPTH(1-34) groups. Histological analyses confirmed these results, emphasizing the stronger bone–implant contact observed in the rhPTH(1-34) group. Furthermore, both PTH analogs significantly promoted bone regeneration in artificially created defects, with the rhPTH(1-34) group displaying a more mature trabecular architecture, as evidenced by a notable increase in the TRAP+ cell count during the bone remodeling assessments.

Revised version: This study investigated the therapeutic effects of rhPTH(1-34) and dimeric R25CPTH(1-34) on bone regeneration and osseointegration in a large animal model with postmenopausal osteoporosis. rhPTH(1-34) and dimeric R25CPTH(1-34) have shown significant clinical efficacy, and although there have been a few studies investigating their effects on bone regeneration in rodents (Garcia et al., 2013), the authors in this study aimed to investigate the effects using a large animal model that more accurately mimics osteoporotic humans (Cortet, 2011). In the evaluation of titanium osseointegration, the rhPTH(1-34) group consistently exhibited enhanced bone mineral density (BMD), bone volume (BV), and other key parameters, thereby indicating superior titanium osseointegration compared to the control and dimeric R25CPTH(1-34) groups (Figure 1E). Histological analyses confirmed these results, emphasizing the stronger bone–implant contact observed in the rhPTH(1-34) group. Furthermore, both PTH analogs significantly promoted bone regeneration in artificially created defects, with the rhPTH(1-34) group displaying a more mature trabecular architecture, as evidenced by a notable increase in the TRAP+ cell count during the bone remodeling assessments (Figure 2M,N).

Added in the revised version: Furthermore, by demonstrating that dimeric R25CPTH(1-34) exhibits a distinct pharmacological profile different from rhPTH(1-34) but still provides a clear anabolic effect in the localized jaw region, the authors have shown that it may possess different potential therapeutic indications from rhPTH(1-34). This biological difference is thought to be due to dimeric R25CPTH(1-34) exhibiting a more preferential binding affinity for the RG versus R0 PTH1R conformation, despite having a diminished affinity for either conformation. (Lee and Lee, 2022; Park et al., 2021) Additionally, the potency of cAMP production in cells was lower for dimeric R25CPTH compared to monomeric R25CPTH, consistent with its lower PTH1R-binding affinity. (Noh et al., 2024) One of the potential clinical advantages of dimeric R25CPTH(1-34) is its partial agonistic effect in pharmacodynamics. This property may allow for a more fine-tuned regulation of bone metabolism, potentially reducing the risk of adverse effects associated with full agonism, such as hypercalcemia and bone resorption by osteolcast activity. Moreover, the dimeric form may offer a more sustained anabolic response, which could be beneficial in the context of long-term treatment strategies. (Noh et al., 2024) Also, the effects of dimer were prominent, as we mentioned better bone formation than the control group.

Removed in the revised version: Furthermore, by demonstrating that dimeric R25CPTH(1-34) exhibits a distinct pharmacological profile different from rhPTH(1-34) but still provides a clear anabolic effect in the localized jaw region, the authors have shown that it may possess different potential therapeutic indications from rhPTH(1-34).

Earlier version: In addition to the effect of systemic bone mineral gain, the unique anabolic feature of PTH has received clinical attention as an emerging strategy due to the increased risk of ONJ following the use of antiresorptive and the rising need for implantation and bone augmentation in the field of orthopedics and maxillofacial surgery (Ruggiero et al., 2022). Moreover, site-specific differential effects of teriparatide have not been clarified, although it represents an issue often associated with selective concentrations of teriparatide that cause anabolic effects on the central skeleton and possible bone mineral decrease on the peripheral skeleton, including the skull (McClung et al., 2005). Moreover, it can be inferred that facial and jaw bones, which have the same developmental origin as the skull through membranous ossification, will show the same bone response as the skull (Setiawati and Rahardjo, 2019). Previous studies have demonstrated that the central skeleton, including the lumbar and thoracic spine and pelvis regions, showed an increase in areal BMD, while the arms, legs, and skull showed a decrease in bone minerals, suggesting an effect of PTH on the redistribution of bone minerals from the peripheral to the central skeleton (Paggiosi et al., 2018).

Revised version: In addition to the effect of systemic bone mineral gain, the unique anabolic feature of PTH has received clinical attention as an emerging strategy due to the increased risk of ONJ following the use of antiresorptive and the rising need for implantation and bone augmentation in the field of orthopedics and maxillofacial surgery (Ruggiero et al., 2022). Moreover, site-specific differential effects of teriparatide have not been clarified, although it represents an issue often associated with selective concentrations of teriparatide that cause anabolic effects on the central skeleton and possible bone mineral decrease on the peripheral skeleton, including the skull (McClung et al., 2005). Moreover, it can be inferred that facial and jaw bones, which have the same developmental origin as the skull through membranous ossification, will show the same bone response as the skull (Setiawati and Rahardjo, 2019). Previous studies have demonstrated that the central skeleton, including the lumbar and thoracic spine and pelvis regions, showed an increase in areal BMD, while the arms, legs, and skull showed a decrease in bone minerals, suggesting an effect of rhPTH(1-34) on the redistribution of bone minerals from the peripheral to the central skeleton (Paggiosi et al., 2018).

Earlier version: However, the results of this study demonstrated that rhPTH(1-34) and dimeric R25CPTH(1-34) significantly improved the osseointegration of bone and titanium, as well as jawbone regeneration. The authors have attributed this phenomenon to the unique anatomical characteristics observed in the jawbone. The jawbone in the human body undergoes the most rapid bone remodeling and has excellent blood flow (Huja et al., 2006). since is continuously exposed to mechanical stress due to mastication and swallowing, thereby suggesting that the net anabolic effect of PTH in the jawbone, which is not part of the central skeleton, is achieved through mechanical loading. A recent study by Robinson et al. demonstrated that PTH and mechanical loading additively stimulate anabolic modeling and synergistically stimulate remodeling in trabecular bone, findings that further support this notion (Robinson et al., 2021). However, further investigation is needed to fully understand this relationship.

Revised version: However, the results of this study demonstrated that rhPTH(1-34) and dimeric R25CPTH(1-34) significantly improved the osseointegration of bone and titanium, as well as jawbone regeneration (Figure 3). The authors have attributed this phenomenon to the unique anatomical characteristics observed in the jawbone. The jawbone in the human body undergoes the most rapid bone remodeling and has excellent blood flow (Huja et al., 2006). since is continuously exposed to mechanical stress due to mastication and swallowing, thereby suggesting that the net anabolic effect of rhPTH(1-34) in the jawbone, which is not part of the central skeleton, is achieved through mechanical loading. A recent study by Robinson et al. demonstrated that rhPTH(1-34) and mechanical loading additively stimulate anabolic modeling and synergistically stimulate remodeling in trabecular bone, findings that further support this notion (Robinson et al., 2021). However, further investigation is needed to fully understand this relationship.

Earlier version: The anabolic effects of PTH have been demonstrated through several large randomized controlled trials (Neer et al., 2001; Tsai et al., 2013), although the drawbacks arising from the unique characteristics of PTH have hindered its broad clinical application. In addition to the disadvantages of continuous self-injection, high cost, and a two-year treatment limit, the greatest concern with prolonged use of teriparatide might be accelerated bone remodeling, subsequent hypercalcemic condition, and increased bone resorption (Burr et al., 2001; Fox, Miller, Newman, et al., 2007; Fox, Miller, Recker, et al., 2007; Neer et al., 2001; Sato et al., 2004; Tsai et al., 2013). However, recent studies have shown that the frequency and dosing of PTH administration can lead to different bone responses. (Yamamoto et al., 2016; Yamane et al., 2017) An in vivo study by Yamamoto et al. reported that lower doses of rhPTH(1-34) with high-frequency administration resulted in the formation of thin trabeculae, osteoclastogenesis, and accelerated bone remodeling, while low-frequency PTH administration showed a phenomenon of modeling-based formation through thicker trabeculae, mature osteoblasts, and new bone formation (Yamane et al., 2017). This supports the notion that modeling-based bone gain can be an important axis in PTH anabolism and the primary principle of PTH as bone remodeling-driven bone anabolism.

Revised version: The anabolic effects of rhPTH(1-34) have been demonstrated through several large randomized controlled trials (Neer et al., 2001; Tsai et al., 2013), although the drawbacks arising from the unique characteristics of rhPTH(1-34) have hindered its broad clinical application. In addition to the disadvantages of continuous self-injection, high cost, and a two-year treatment limit, the greatest concern with prolonged use of teriparatide might be accelerated bone remodeling, subsequent hypercalcemic condition, and increased bone resorption (Burr et al., 2001; Fox, Miller, Newman, et al., 2007; Fox, Miller, Recker, et al., 2007; Neer et al., 2001; Sato et al., 2004; Tsai et al., 2013). However, recent studies have shown that the frequency and dosing of rhPTH(1-34) administration can lead to different bone responses. (Yamamoto et al., 2016; Yamane et al., 2017) An in vivo study by Yamamoto et al. reported that lower doses of rhPTH(1-34) with high-frequency administration resulted in the formation of thin trabeculae, osteoclastogenesis, and accelerated bone remodeling, while low-frequency rhPTH(1-34) administration showed a phenomenon of modeling-based formation through thicker trabeculae, mature osteoblasts, and new bone formation (Yamane et al., 2017). This supports the notion that modeling-based bone gain can be an important axis in PTH anabolism and the primary principle of PTH as bone remodeling-driven bone anabolism.

Added in the revised version: The use of daily injections in this study was intended to simulate intermittent PTH therapy, a well-established clinical approach for managing osteoporosis and enhancing bone regeneration. Intermittent administration of PTH, as opposed to continuous exposure, is critical for maximizing the anabolic response while minimizing the catabolic effects that are associated with higher frequency or continuous hormone levels. Our findings support the notion that even with daily administration, both rhPTH(1-34) and dimeric dimeric R25CPTH(1-34) promote bone formation and osseointegration, consistent with the outcomes expected from intermittent therapy. It’s important for future research to consider the dosage and timing of administration to further optimize the therapeutic benefits of PTH analogs (Dempster et al., 2001; Hodsman et al., 2005).

The limitation of this study is that the therapeutic responses of rhPTH(1-34) and dimeric R25CPTH(1-34) were focused on local surgical interventions, meaning that we could not investigate the central skeletal responses, such as in the femur and lumbar. Therefore, further research is needed to investigate these different responses by administering the PTH analogs at various frequencies and doses.

Overall, the study demonstrated the therapeutic effects of rhPTH(1-34) and dimeric R25CPTH(1-34) on bone regeneration and titanium osseointegration using a beagle model with osteoporosis. Validation of the anabolic effects of rhPTH(1-34) and dimeric R25CPTH(1-34) in large animals has resulted in a broader understanding of their physiological and therapeutic functions and further expands their potential applications.

Material and Methods

Animal preparation

This study was conducted in compliance with the ARRIVE guidelines and approved by the Animal Research Committee of Cronex Co., Ltd., Hwaseong, South Korea. (CRONEX-IACUC 201801002). All animal experiments, including animal selection, management, preparation, and surgical protocols were conducted in compliance with the Ewha Womans University rules for animal experiments. The animals were housed in a standard laboratory environment (21 ± 1 ◦C with 40–70% humidity; 12Lhours light/dark cycle) with a standardized food/water supply.

Beagles were chosen for this study because the bone size and dentition could accommodate human dental implants and the application of mechanical force to implants. To induce osteoporosis, 12 female beagles underwent bilateral ovariectomy (OVX) at 12 weeks of age, followed by osteoporosis development for 12 weeks before being used for the next experiments at the age of 24 weeks. Health and oral hygiene were checked and maintained daily (Figure 1A).

Experiment protocol

Anesthesia was induced using Zolazepam/Tiletamine (10 mg/kg body weight, Zoletil; Virbac Laboratories, Carros, France) and xylazine hydrochloride (Rumpun®, Bayer, Leverkusen, Germany), by intramuscular injection. The beagles were anesthetized using inhalation anesthesia for the implant surgery, and antibiotics (Ceftriaxone, Kyungdong Pharm, Seoul, South Korea) were administered for three days.

Both mandibular premolars 1–4 were extracted at 12 weeks after OVX. Then, 12 weeks later, three dental titanium implants (TS III® 3.0 x 10 mm, Osstem, Seoul, Korea) were inserted into the right lower jaw of each animal via a conventional implant surgical procedure under saline irrigation. Each implant was placed over at least a 3.0 mm distance. In detail, the first implant was placed over 3.0 mm behind the canine. The second and third implants were inserted to create a 3 mm circumferential bony defect using a Ø 6 mm trephine bur. Then, an additional bone graft (Bio-oss® small particle 1.0 g and Bio-gide®, Geistlich, Switzerland) was placed on the bone defect around the 3rd implant, and the bone gap remained with nothing at the 2nd implant. On the left lower jaw, two artificial bone defects of 5 x 10 mm were made using trephine bur. The anterior hole remained defective without a graft, while a bone graft (Bio-oss® small particle 1.0 g and Bio-gide®, Geistlich, Switzerland) was applied to the posterior hole.

Earlier version: After 2 weeks of healing, twelve female OVX beagles were randomly designated into three groups as follows: 1) control group with normal saline injection, 2) PTH(1-34) group with daily 40 μg/day rhPTH(1-34) injection (Forsteo®, Eli Lilly and Company, Indianapolis, IN, USA), and 3) dimeric R25CPTH(1-34) group with 40 μg/day injection (dimeric R25CPTH(1-34), chemically synthesized by the by Anygen, Gwangju, Republic of Korea). Animals were injected subcutaneously for 10 weeks, after which, they were euthanized and the bone regeneration and implant osseointegration were evaluated (Figure 1A, B, and C). To analyze the dynamic bone formation, 10 mg/kg of calcein green (Sigma, St, Louis, USA) and 30 mg/kg of oxytetracycline yellow (Fluka, Shanghai, China) were intramuscularly administered at 2 weeks and 10 weeks after operation.

Revised version: After 2 weeks of healing, twelve female OVX beagles were randomly designated into three groups as follows: 1) control group with normal saline injection, 2) PTH(1-34) group with daily 40 μg/day rhPTH(1-34) injection (Forsteo®, Eli Lilly and Company, Indianapolis, IN, USA), and 3) dimeric R25CPTH(1-34) group with 40 μg/day injection (dimeric R25CPTH(1-34), chemically synthesized by the by Anygen, Gwangju, Republic of Korea). Each animal received one injection per day, aimed at replicating the intermittent rhPTH(1-34) exposure proven beneficial for bone regeneration and overall skeletal health in clinical settings (Neer et al., 2001; Kendler et al., 2018). This regimen was chosen to investigate the potential anabolic effects of these specific PTH analogs under conditions closely resembling therapeutic use. Animals were injected subcutaneously for 10 weeks, after which, they were euthanized and the bone regeneration and implant osseointegration were evaluated (Figure 1A, B, and C). To analyze the dynamic bone formation, 10 mg/kg of calcein green (Sigma, St, Louis, USA) and 30 mg/kg of oxytetracycline yellow (Fluka, Shanghai, China) were intramuscularly administered at 2 weeks and 10 weeks after operation.

Micro-computed tomographic (µCT) analysis

Earlier version: To examine the microarchitectural effects of PTH(1-34) and (Cys25)PTH(1-84) on bone regeneration and osseointegration, radiographic analysis was performed using µCT on the right mandible (Figure 1D). The specimens were fixed in 4% paraformaldehyde for 48 hours before being assessed by micro-computed tomography (µCT, SkyScan1173 ver. 1.6, Bruker-CT, Kontich, Belgium). The specimens were imaged with a pixel size of 29.83 µm. The voltage and current intensities of the images were 130 kV and 60 µA, respectively. The regions of interest (ROI) were determined as the 10 x 10 mm square area, located 3 mm from the bottom of the implant. Bone mineral density (BMD), bone volume (BV; mm3), trabecular number (Tb.N; 1/mm), trabecular thickness (Tb.Th; µm), and trabecular separation (Tb.S; µm) were analyzed.

Revised version: To examine the microarchitectural effects of PTH(1-34) and dimeric R25CPTH(1-34) on bone regeneration and osseointegration, radiographic analysis was performed using µCT on the right mandible (Figure 1D). The specimens were fixed in 4% paraformaldehyde for 48 hours before being assessed by micro-computed tomography (µCT, SkyScan1173 ver. 1.6, Bruker-CT, Kontich, Belgium). The specimens were imaged with a pixel size of 29.83 µm. The voltage and current intensities of the images were 130 kV and 60 µA, respectively. The regions of interest (ROI) were determined as the 10 x 10 mm square area, located 3 mm from the bottom of the implant. Bone mineral density (BMD), bone volume (BV; mm3), trabecular number (Tb.N; 1/mm), trabecular thickness (Tb.Th; µm), and trabecular separation (Tb.S; µm) were analyzed.

Histological and histomorphometric analysis

Bone histomorphometric parameters were computed and shown in accordance with recommendations by the ASBMR histomorphometric nomenclature committee. (Dempster et al., 2013) Goldner’s trichrome and Masson trichrome staining were performed on both the right implantation and left bone defect sites, respectively. Specimens were dehydrated in increasing concentrations of ethanol and embedded in a mixture of ethanol and Technovit 7200 resin (Heraeus Kulzer, Wehrheimm, Germany), with an increasing ratio of resin. Following resin infiltration, the specimens were hardened in a UV embedding system (KULZER EXAKT 520, Norderstedt, Germany) for a day. The undecalcified specimens were cut using an EXAKT diamond cutting system (EXAKT 300 CP, Norderstedt, Germany), and the soft tissue and bone were attached to an acryl slide by an adhesive system. The section width of the specimen was adjusted to 40 ± 5µm using a grinding system (EXAKT 400CS, KULZER, Norderstedt, Germany). The specimens on the right implantation site were stained with Goldner’s trichrome and photographed by a Panoramic 250 Flash III system (3DHISTECH Ltd., Budapest, Hungary). The bone–implant contact ratio (BIC, %) was assessed as the linear percentage of the interface with direct contact between the bone and implant to the total interface of the implant using CaseViewer program software (3DHISTECH Ltd.).

Tartrate-resistant acid phosphatase (TRAP) assay was performed on the left bone defect sites. The bone specimens were fixed in 4% paraformaldehyde overnight and decalcified in 10% ethylenediaminetetraacetic acid (EDTA) for 7 days. The decalcifying solution was changed every other day. The specimens were embedded in paraffin and cut into sections. According to the manufacturer’s instructions, the sections were deparaffinized and stained using a TRAP staining kit (Sigma, St. Louis. MO, USA). The number of TRAP-positive cells in the sections was counted under a microscope (DM2500, Leica Microsystems, Wetzlar, Germany).

Serum biochemical analysis

Earlier version: Fasting blood samples were drawn in the morning at baseline (T0; 12 weeks), the start of the injection (T1), and at euthanasia (T2) for P1NP (procollagen type I N-terminal propeptide), PTH (parathyroid hormone), CTx (C-terminal telopeptide), and calcium, and phosphorus. Calcium and phosphorus were analyzed using the Beckman AU480 Chemistry Analyzer (Beckman Coulter AU480). P1NP (procollagen type I N-terminal propeptide ELISA kit, Mybiosource), PTH (parathyroid hormone ELISA kit, Aviva Systems Biology), and CTX (C-terminal telopeptide (CTx-I) ELISA kit, Mybiosource) were analyzed by the ELISA method, according to the ELISA kit manufacturer’s instructions. In all analyses, the measured values were below the limit of quantification for the standard curve.

Revised version: Fasting blood samples were drawn in the morning at baseline (T0; 12 weeks), the start of the injection at 26th weeks(T1), and euthanasia at 36th weeks (T2) for P1NP (procollagen type I N-terminal propeptide), PTH (parathyroid hormone), CTx (C-terminal telopeptide), and calcium, and phosphorus (Figure 5A). Calcium and phosphorus were analyzed using the Beckman AU480 Chemistry Analyzer (Beckman Coulter AU480). P1NP (procollagen type I N-terminal propeptide ELISA kit, Mybiosource), PTH (parathyroid hormone ELISA kit, Aviva Systems Biology), and CTX (C-terminal telopeptide (CTx-I) ELISA kit, Mybiosource) were analyzed by the ELISA method, according to the ELISA kit manufacturer’s instructions. In all analyses, the measured values were below the limit of quantification for the standard curve.

Statistical analysis

Data for microarchitectural, histomorphometric, and serum biochemical analyses were expressed as mean and standard deviation (SD). Non-parametric tests were performed, including the Mann– Whitney and Kruskal–Wallis tests. Group differences in serum markers over time were compared by repeatedly measuring the analysis of variance. Statistical analysis was performed using SPSS 26 (IBM Corp., USA) and Prism 10 (GraphPad, San Diego, CA, USA). P values of < 0.05 were set as statistically significant.

Added in the revised version: Supporting information